Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Cynthia Bachor"
COLLABORATORS = "2111599"
```

---

---
# Welcome to CE880
### This is your week-7 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 

---
## Question 1: 

#### Bootstrapping

We are generating a random vector, which will be our population, whose law of distribution (expected to be the same as that of the population) is unknown. Write a function to pick a sample from our population and by applying the Bootstrapping procedure, where sample `size=300`: and `B=1000` number of re-sampling. Calculate the sample mean from 1000 bootstraping samples and compare with population mean. 

<!-- grade_id: my_bootstrapping_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
np.random.seed(123)
pop = np.random.randint(0,500 , size=1000)
print(pop.mean()) # Print population mean
sample = np.random.choice(pop, size=300) #so n=300

def my_bootstrapping():
    """Write a function to pick a sample from our population and by applying the
    Bootstrapping procedure, where sample size=300: and B=1000 number of re-sampling """
    mean_sample =[]
    for a in range(1):
        mean_sample.append(sample.mean())
    return np.mean(mean_sample)
    raise NotImplementedError()
```

Output:
```
253.241
```

<!-- grade_id: my_bootstrapping_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_bootstrapping(), 253.241, rel_tol = 0.05)
```

---
## Question 2: 

#### Majority Voting
Given a list = `[100,200,300,100,200,100,100,10,30,20,20,10]`. Write a function to find the majority voting.

<!-- grade_id: my_pairedttest_ans, points: 0, grade: false, solution: true -->
```python
a = np.array([100,200,300,100,200,100,100,10,30,20,20,10])

def my_majorityvote():
    """Write a function to find the majority voting."""
    # YOUR CODE HERE
    import math
    from collections import Counter
    c = Counter(a)
    start_counter = 0
    num = a[0]

    for i, vote_count in c.items():
        if(vote_count > start_counter):
          start_counter = vote_count
          num = i
    
    return num
print(my_majorityvote())
```

Output:
```
100
```

<!-- grade_id: my_pairedttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_majorityvote(), 100, rel_tol = 0.05)
```

---
## Question 3: 

#### Decision Tree
USe the `IRIS` dataset from `sklearn` to build a decision tree classifier. The dataset contains information for three classes of the IRIS plant, namely `IRIS Setosa`, `IRIS Versicolour`, and `IRIS Virginica`, with the following attributes: `sepal length`, `sepal width`, `petal length`, and `petal width`. Your task is to train a decision tree on `training` data and return the `test` accuracy from `[0-100]`. Hint: multiply for `accuracy_score` it with 100. 

<!-- grade_id: my_DTree_ans, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from sklearn.datasets import load_iris
from sklearn.metrics import accuracy_score
from sklearn.tree import DecisionTreeClassifier

# Reading the Iris.csv file
data = load_iris()

# Extracting Attributes / Features
X = data.data

# Extracting Target / Class Labels
y = data.target

# Import Library for splitting data
from sklearn.model_selection import train_test_split

# Creating Train and Test datasets
X_train, X_test, y_train, y_test = train_test_split(X,y, random_state = 50, test_size = 0.25)

def my_DTree():
    """Write code to to train a decision tree classifier and return the test accuracy"""
    # YOUR CODE HERE
    from sklearn.model_selection import train_test_split
    dec_tree = DecisionTreeClassifier() #assigning the classifier to train the model
    dec_tree.fit(X_train, y_train)
    Check_model_accuracy = accuracy_score(y_test, dec_tree.predict(X_test)) #predecting and checking accuracy score
    Check_model_accuracy = Check_model_accuracy * 100
    return Check_model_accuracy

    raise NotImplementedError()
```

<!-- grade_id: my_DTree_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_DTree(), 94.736, rel_tol = 0.05)
```

```python

```